```python
import collections
import string
import itertools
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

```python
def primes():
    # keeps track of the current composite we are checking
    # composites has all primes encountered so far as values
    # the keys are the next composite number they divide that we haven't reached
    # composites[Composite number] = [primes whose multiples we're tracking]
    composites = {}
    q = 2 #running integer checked for primeness

    while True: #make infinite loop
        if q not in composites: # q is prime
            composites[q*q] = [q] #start tracking multiple of q
            yield q
        else: #q is composite
            for pr in composites[q]:
                # this marks the next multiples of the prime numbers we are tracking
                # does pr + q exist as a key? if it doesn't assign [] then append p
                composites.setdefault(pr+q, []).append(pr) 
            del composites[q]

        q += 1
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

```python
itertools.repeat(5,2)
```

Output:
```
repeat(5, 2)
```

```python
def ruler():
    num = 1
    count = 1

    while True: # infinite generator
        for i in range(count): # number of repeats
            for j in range(count): #number of new numbers that repeat
                yield num
            num += 1
        count += 1
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

```python
# count = 1 # count
# curr_term = "1211" # the current term

# digit = curr_term[0] # the digit in the run you are counting
# next_term = ""

# for char in range(len(curr_term)-1):

#     if curr_term[char+1] == digit: # next char equals current
#         count += 1
#     else:
#         run_term = str(count) + digit
#         next_term += run_term
#         digit = curr_term[char+1]
#         count = 1
# next_term += str(count) + digit

# next_term
```

Output:
```
'111221'
```

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
def look_say():
    curr_term = "1"

    while True:
        yield curr_term

        count = 1
        digit = curr_term[0] # the digit in the run you are counting
        next_term = ""

        for char in range(len(curr_term)-1):

            if curr_term[char+1] == digit: # next char equals current
                count += 1
            else:
                next_term += str(count) + digit
                digit = curr_term[char+1]
                count = 1
        next_term += str(count) + digit
        curr_term = next_term
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
    alph = "abcdefghijklmnopqrstuvwxyz"
    enc_s = ""

    for char in s:
        reg_ind = alph.index(char)
        enc_s += alph[(reg_ind + k) % 26]
    return enc_s
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    alph = "abcdefghijklmnopqrstuvwxyz"
    dec_s = ""

    for char in s:
        ciph_ind = alph.index(char)
        dec_s += alph[(ciph_ind - k) % 26]
    return dec_s
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


```python
# test = "([)]"

# stack = []
# matches = {"(": ")", "[": "]", "{": "}", ")": "(", "]": "[", "}": "{"}
# curr = test[0]
# valid = True

# for curr in test:

#     if curr in "({[": # curr is an opener
#         stack.append(curr) 
#         print("add opener to stack")
#     else: # curr is a closer
#         if len(stack) == 0:
#             print("False: closing bracket has no opener")
#             valid = False
#             break
#         opener = stack.pop() #remove opener
#         if matches[curr] != opener:
#             print("False")
#             valid = False
#             break
#         print("bracket matches. stack")

# if valid:
#     if len(stack) == 0:
#         print("True")
#     else:
#         print("False")
# # print("True")

# valid
```

<!-- grade_id: is_valid, points: 0, grade: false, solution: true -->
```python
def is_valid(s):
    stack = []
    matches = {"(": ")", "[": "]", "{": "}", ")": "(", "]": "[", "}": "{"}

    for curr in s:

        if curr in "({[": # curr is an opener
            stack.append(curr) 
            # print("add opener to stack")
        else: # curr is a closer
            if len(stack) == 0:
                print("False: closing bracket has no opener")
                return False
            opener = stack.pop() #remove opener
            if matches[curr] != opener:
                print("False: next opener does not match")
                return False


    if len(stack) == 0:
        return True
    else:
        return False
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```

Output:
```
False: next opener does not match
```